```python
# Initialize Otter
import otter
grader = otter.Notebook("fifa_21_raw_data.ipynb")
```

## Week 2 - FIFA '21 Player Ratings

This notebook is created for the week 2 assignment. Here, you have a very messy and raw dataset of EA Sports' installment of their hit FIFA series - FIFA21, which was scraped from sofifa.com.

Your task in this assignment is to clean up this dataset using the `Hints` provided in the ReadMe file

```python
import pandas as pd
import numpy as np
import os
%matplotlib inline
import otter
grader = otter.Notebook()
```

```python
pd.set_option('display.max_rows', None)
pd.set_option('display.max_columns', None)
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

```python
fifa_dataset = pd.read_csv('./fifa_21_raw_data.csv')

def check_data(fifa_dataset):
    if fifa_dataset.empty:
        return True
    else:
        return False
    
check_data(fifa_dataset)
```

Output:
```
C:\Users\mercy\AppData\Local\Temp\ipykernel_17276\3062757269.py:1: DtypeWarning: Columns (76) have mixed types. Specify dtype option on import or set low_memory=False.
  fifa_dataset = pd.read_csv('./fifa_21_raw_data.csv')
False
```

```python
fifa_dataset.head()
```

Output:
```
ID               Name                      LongName  \
0  158023           L. Messi                  Lionel Messi   
1   20801  Cristiano Ronaldo  C. Ronaldo dos Santos Aveiro   
2  200389           J. Oblak                     Jan Oblak   
3  192985       K. De Bruyne               Kevin De Bruyne   
4  190871          Neymar Jr    Neymar da Silva Santos Jr.   

                                           photoUrl  \
0  https://cdn.sofifa.com/players/158/023/21_60.png   
1  https://cdn.sofifa.com/players/020/801/21_60.png   
2  https://cdn.sofifa.com/players/200/389/21_60.png   
3  https://cdn.sofifa.com/players/192/985/21_60.png   
4  https://cdn.sofifa.com/players/190/871/21_60.png   

                                           playerUrl Nationality  Age  ↓OVA  \
0  http://sofifa.com/player/158023/lionel-messi/2...   Argentina   33    93   
1  http://sofifa.com/player/20801/c-ronaldo-dos-s...    Portugal   35    92   
2  http://sofifa.com/player/200389/jan-oblak/210006/    Slovenia  
```

```python
grader.check("q1")
```

Output:
```
File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\notebook.py:255, in Notebook.check(self, question, global_env)
    242 """
    243 Runs tests for a specific question against a global environment. If no global environment
    244 is provided, the test is run against the calling frame's environment.
   (...)
    252     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    253 """
    254 self._logger.info(f"Running check for question: {question}")
--> 255 test_path, test_name = resolve_test_info(
    256     self._tests_dir,
    257     self._notebook,
    258     self._tests_url_prefix,
    259     question,
    260 )
    262 self._logger.debug(f"Resolved test path: {test_path}")
    263 self._logger.debug(f"Resolved test name: {test_name}")

File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\utils.py:327, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    325 el
```

**Question 2:** Remove all duplicate player `name` from the dataset. Remember, you need to leave the first occurence and remove other ones.

```python
def remove_duplicate(dataset):
    dataset.drop_duplicates(subset=['Name'], inplace=True)
```

```python
grader.check("q2")
```

Output:
```
File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\notebook.py:255, in Notebook.check(self, question, global_env)
    242 """
    243 Runs tests for a specific question against a global environment. If no global environment
    244 is provided, the test is run against the calling frame's environment.
   (...)
    252     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    253 """
    254 self._logger.info(f"Running check for question: {question}")
--> 255 test_path, test_name = resolve_test_info(
    256     self._tests_dir,
    257     self._notebook,
    258     self._tests_url_prefix,
    259     question,
    260 )
    262 self._logger.debug(f"Resolved test path: {test_path}")
    263 self._logger.debug(f"Resolved test name: {test_name}")

File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\utils.py:327, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    325 el
```

**Question 3:** Remove all leading white spaces or newline "\n" in the `club` column. For example, `"\n\n\n\nFC Barcelona"` should be transformed into `"FC Barcelona"`



```python
def remove_newline(dataset):
    dataset['Club'] = dataset['Club'].str.lstrip('\n ')
    return dataset
```

```python
grader.check("q3")
```

Output:
```
File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\notebook.py:255, in Notebook.check(self, question, global_env)
    242 """
    243 Runs tests for a specific question against a global environment. If no global environment
    244 is provided, the test is run against the calling frame's environment.
   (...)
    252     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    253 """
    254 self._logger.info(f"Running check for question: {question}")
--> 255 test_path, test_name = resolve_test_info(
    256     self._tests_dir,
    257     self._notebook,
    258     self._tests_url_prefix,
    259     question,
    260 )
    262 self._logger.debug(f"Resolved test path: {test_path}")
    263 self._logger.debug(f"Resolved test name: {test_name}")

File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\utils.py:327, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    325 el
```

**Question 4:** List the TOP 5 countries with the highest number of representation of players in the data. Create a dictionary called `top_5_country` using the country names as the key and the number of players as the value. For example,

`top_5_country` = {'Kenya': 15,
                  'Nigeria': 45,
                  'Ghana': 14,
                  'Canada': 20,
                  'USA': 88
                 }


```python
top_5_country = fifa_dataset['Nationality'].value_counts().head(5).to_dict()

```

```python
grader.check("q4")
```

Output:
```
File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\notebook.py:255, in Notebook.check(self, question, global_env)
    242 """
    243 Runs tests for a specific question against a global environment. If no global environment
    244 is provided, the test is run against the calling frame's environment.
   (...)
    252     ``otter.test_files.abstract_test.TestFile``: the grade for the question
    253 """
    254 self._logger.info(f"Running check for question: {question}")
--> 255 test_path, test_name = resolve_test_info(
    256     self._tests_dir,
    257     self._notebook,
    258     self._tests_url_prefix,
    259     question,
    260 )
    262 self._logger.debug(f"Resolved test path: {test_path}")
    263 self._logger.debug(f"Resolved test name: {test_name}")

File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\utils.py:327, in resolve_test_info(tests_dir, nb_path, tests_url_prefix, question)
    325 el
```

<!-- BEGIN QUESTION -->

**Question 5:** What insight can you derive from this data?

From the data, the average football age is 25, mean height is approximately 181.28 cm, mean weight is approximately 75.08 kg. Argentina has the highest number of skilled players.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\notebook.py:427, in Notebook.export(self, nb_path, export_path, pdf, filtering, pagebreaks, files, display_link, force_save, run_tests)
    423 self._log_event(EventType.BEGIN_EXPORT)
    425 files = [] if files is None else files
--> 427 nb_path = self._resolve_nb_path(nb_path)
    428 self._logger.debug(f"Resolved notebook path: {nb_path}")
    430 if force_save:

File c:\Users\mercy\AppData\Local\Programs\Python\Python311\Lib\site-packages\otter\check\notebook.py:236, in Notebook._resolve_nb_path(self, nb_path, fail_silently)
    233         nb_path = notebooks[0]
    235 if nb_path is None and not fail_silently:
--> 236     raise ValueError("Could not resolve notebook path")
    238 return nb_path

ValueError: Could not resolve notebook path
```

Tests for q1 (points: 1)
```
>>> fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', engine='python')
>>> 
>>> check_data(fifa_dataset)
False
```


Tests for q2 (points: 1)
```
>>> remove_duplicate(fifa_dataset)
>>> fifa_dataset['Name'].duplicated().any() == False
True
```


Tests for q3 (points: 1)
```
>>> len(fifa_dataset["Name"].unique()) > 1000
True
```
```
>>> type(fifa_dataset["Club"]) == pd.Series
True
```


Tests for q4 (points: 1)
```
>>> type(top_5_country) is dict
True
```
```
>>> len(top_5_country) == 5
True
```
